Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Artem Burtsev"
COLLABORATORS = "Me and only"
```

---

# Lagrange interpolating polynomial.

Implement a class which constructs the Lagrange interpolating polynomial given data arrays `xk` and `yk`.

<!-- grade_id: cell-d459dee4d09527c4, points: 0, grade: false, solution: true -->
```python
import numpy as np

class LagrangeInterpolator:
    """Lagrange interpolating polynomial.
    
    Given a set of pairs ``(x_k, y_k)``, construct 
    a Lagrange polynomial ``f(x)``, such that
    
    .. math::

        f(x_k) = y_k   for k =0, ..., n-1
    
    Parameters
    ----------
    xk : array_like, shape(n,)
        Abscissas
    yk : array_like, shape(n,)
        Ordinates
    
    Attributes
    ----------
    __call__
    
    """
    def __init__(self, xk, yk):
        self.xk = np.asarray(xk, dtype=float)
        self.yk = np.asarray(yk, dtype=float)
        
    def __call__(self, x):
        """Evaluate the interpolator at a given point.
        
        Parameters
        ----------
        x : float
        
        Returns
        -------
        the value of the interpolator at ``x``.
        """
        n = len(self.xk)
        
        def l(x, index):
            result = 1
            for current_index in range(n):
                if current_index != index:
                    result *= (x - self.xk[current_index])/(self.xk[index] - self.xk[current_index])
            return result
        
        value = 0
        for index in range(n):
            value += self.yk[index] * l(x, index)
        return value
```

<!-- grade_id: cell-02e34e05d317dfb2, points: 7, grade: true, solution: false -->
```python
def runge_func(x, a=25):
    return 1.0 / (1.0 + a*x**2)

xx = np.linspace(-2, 2, 21)
yy = runge_func(xx)

lagr = LagrangeInterpolator(xx, yy)

from numpy.testing import assert_allclose

assert_allclose(yy,
                [lagr(xval) for xval in xx],
                atol=1e-14)
```

## The Runge phenomenon

Now consider the Runge function, $y = 1/(1 + 25x^2)$. Interpolate this function on the interval $x\in [-2, 2]$, 
using Lagrange polynomials with $m=3, 5, 7, 11$. Use a uniform mesh. Plot the resulting interpolants together with the original function, $f(x)$.
Discuss.

(extra exercise, ungraded)

```python
# YOUR CODE AND COMMENTS HERE
```

Implement a function which returns Chebyshev nodes, i.e. the roots of the Chebyshev polynomial of the first kind  of a given degree.

<!-- grade_id: cell-43b1fb4f14c698bc, points: 0, grade: false, solution: true -->
```python
def cheb_nodes(n, a=-1, b=1):
    r"""Chebyshev nodes of degree $n$ on $[a, b]$
    """
    nodes = [1/2*(a + b) + 1/2*(b - a) * np.cos((2*k + 1)*np.pi/(2*n)) for k in range(n)]
    return nodes[::-1]
```

<!-- grade_id: cell-1ea45f1e8dc8a36d, points: 3, grade: true, solution: false -->
```python
nodes_11 = cheb_nodes(11)
nodes_11 = np.asarray(nodes_11)
assert (nodes_11[1:] > nodes_11[:-1]).all()

from scipy.special import roots_chebyt
nodes, weights = roots_chebyt(5) 

assert_allclose(cheb_nodes(5),
                nodes, atol=1e-14)

assert_allclose(cheb_nodes(5, a=-1, b=3),
                nodes*2 + 1, atol=1e-14)
```

Repeat the Lagrange interpolation of the Runge function using Chebyshev nodes. Plot the interpolants.
Also plot the cubic spline interpolation of the same data (`scipy.interpolate.CubicSpline`).
Compare the magnitude of the Runge phenomenon for the uniform mesh and the Chebyshev mesh. Does the spline interpolant demonstrate the Runge phenomenon?

(extra exercise, ungraded)

```python
# YOUR CODE AND COMMENTS HERE
```