Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Artem Burtsev"
COLLABORATORS = "Me and only"
```

---

```python
import numpy as np
```

# Newton's method for systems of nonlinear equations

Implement the function which performs Newton iteration for a system of nonlinear equations, $\vec{f}(x) = 0$, with a known Jacobi matrix, $J(x)$. Your functions must find a root of a vector function $\vec{f}(x)$ with a predefined tolerance $\epsilon$ (i.e., iterations stop when the norm of the difference between succesive approximations is smaller then $\epsilon$).

<!-- grade_id: cell-4580b8353412b2d3, points: 0, grade: false, solution: true -->
```python
def newton_system(F, jcb, X0, eps=1e-5, maxiter=100):
    """Newton iteration for the system $F(x) = 0$.
    
    Parameters
    ----------
    F : callable
        Vector-function of the l.h.s. of the system, $F(x)$. The function `F` receives
        a vector `x` of length `n` and returns a vector of length `n`.
    jcb : callable
        Jacobian of `f`.
    X0 : array-like of floats, shape (n,)
        The starting point for the iteration
    eps : float
        The stopping criterion is that iterations stop when the norm of the difference
        between successive iterates is smaller then `eps`.
        Default is 1e-5
    maxiter : int
        Maximum allowed number of iterations.
        Default is 100.
        This parameter is needed to avoid infinite loops.
    
    Returns
    -------
    X : array-like of floats, shape (n,)
        The approximation to the root of $F(x) = 0$
    niter : int
        The number of iterations.
    """
    X = X0
    X_calculated = 0
    condition = 1
    niter = 0
    while condition >= eps:
        jacobian_matrix = jcb(X)
        X_calculated = -np.linalg.solve(jacobian_matrix, np.eye(jacobian_matrix.shape[0])) @ F(X) + X
        condition = np.linalg.norm(X_calculated - X)
        X = X_calculated
        
        if niter == maxiter:
            return X_calculated, maxiter
        
        niter += 1
    return X_calculated, niter
```

<!-- grade_id: cell-32a522217549fd0c, points: 3, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

def func(X):
    return np.array([X[0]**2 - 2*X[1]**4 + 1,
                     X[0] - X[1]**3 + 1])

def jac(X):
    return np.array([[2.*X[0], -8.*X[1]**3],
                     [1., -3.*X[1]**2]])

xx, nit = newton_system(func, jac, np.array([3., 2.]))
assert_allclose(xx, np.array([4.02879487, 1.71325223]), atol=1e-5)
assert nit < 10

x0 = np.array([1., 2.])
xx, nit = newton_system(func, jac, x0)
assert_allclose(xx, np.asarray([-0.35119943,  0.86570596]), atol=1e-5)
assert nit < 10

# also check that the initial value is not modified by the function
assert_allclose(x0, np.array([1., 2.]), atol=1e-15)
```

Write a function which returns the Jacobian matrix of the function of three variables, `func2`:

<!-- grade_id: cell-d091fffbc05825d7, points: 0, grade: false, solution: true -->
```python
def func2(X):
    x, y, z = X
    return np.array([x**2 + y**2 + z**2 - 25,
                     x*y + y*z + z*x - 5,
                     x + y - 3])

def jac2(X):
    r"""Return the Jacobian of `func2(X)`.
    
    
    The Jacobian matrix is defined as 
    
    $$
    J_{ij} = \partial f_i / \partial x_j
    $$
    
    so that the first row contains the derivatives of $f_0$ with respect to the
    first, second etc coordinates; the second row contains the derivatives of
    $f_1$ with respect to the first, second etc coordinates; and so on.
    """
    x, y, z = X
    return np.array([[2*x, 2*y, 2*z],
                    [y+z, x+z, x+y],
                    [1,1,0]])
```

<!-- grade_id: cell-3a1ae706ab35a030, points: 3, grade: true, solution: false -->
```python
X = [1, 2, 3]
assert jac2(X).shape == (3, 3)

```

<!-- grade_id: cell-7f2820844f3069af, points: 4, grade: true, solution: false -->
```python
# more testing in this cell, keep it intact

```

```python

```