Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Artem Burtsev"
COLLABORATORS = "Me and only"
```

---

```python
import matplotlib.pyplot as plt
import numpy as np
```

# I. Root-finding: Newton iteration

Write a function which performs Newton's iteration for a given function  $f(x)$  with known derivative  $f'(x)$. Your function should find the root of  $f(x)$  with a predefined absolute accuracy  $\epsilon$.

<!-- grade_id: cell-e47469fedb8cc24a, points: 0, grade: false, solution: true -->
```python
def newton_iteration(f, fder, x0, eps=1e-5, maxiter=100):
    """Find a root of $f(x) = 0$ via Newton's iteration starting from x0.
    
    Parameters
    ----------
    f : callable
        The function to find a root of.
    fder : callable
        The derivative of `f`.
    x0 : float
        Initial value for the Newton's iteration.
    eps : float
        The target accuracy. 
        The iteration stops when the distance between successive iterates is below `eps`.
        Default is 1e-5.
    maxiter : int
        The maximum number of iterations (default is 1000.)
        Iterations terminate if the number of iterations exceeds `maxiter`.
        This parameter is only needed to avoid infinite loops if iterations wander off.
        
    Returns
    -------
    x : float
        The estimate for the root.
    niter : int
        The number of iterations.
    """
    next_x = 0
    x = x0
    
    condition = 1
    niter = 0
    while condition >= eps:
        next_x = x - f(x) / fder(x)
        condition = np.linalg.norm(next_x - x)
        x = next_x
        if niter == maxiter:
            return next_x, maxiter
        niter += 1
    return next_x, niter
```

<!-- grade_id: cell-0c23dfbb8447f5aa, points: 5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

xx, nit = newton_iteration(lambda x: x**2 - 1, lambda x: 2.*x, x0=4)
assert_allclose(xx, 1.0, atol=1e-5)
assert nit < 10

xx, nit = newton_iteration(lambda x: x**2 - 1, lambda x: 2.*x, x0=-4)
assert_allclose(xx, -1.0, atol=1e-5)
assert nit < 10


from math import log, exp

xx, nit = newton_iteration(lambda x: exp(x) - 2, lambda x: exp(x), x0=4, eps=1e-10)
assert_allclose(xx, log(2), atol=1e-10)
```

# II. Multiple roots. Modified Newton iteration.

Consider a function which has a multiple root (for instance, $f(x) = (x^2 - 1)^2$). Implement the modified Newton iteration,

$$
x_{n+1} = x_{n} - m \frac{f(x_n)}{f'(x_n)}
$$

where $m$ is a parameter.

<!-- grade_id: cell-2f641d195eae18f0, points: 0, grade: false, solution: true -->
```python
def mod_newton(f, fder, x0, m, eps=1e-5, maxiter=100):
    r"""Find a root of $f(x) = 0$ via Newton's iteration starting from x0.
    
    Parameters
    ----------
    f : callable
        The function to find a root of.
    fder : callable
        The derivative of `f`.
    x0 : float
        Initial value for the Newton's iteration.
    m : int
        The iteration is $x_{n+1} = x_n - m f(x_n) / f'(x_n)`
    eps : float
        The target accuracy. 
        The iteration stops when the distance between successive iterates is below `eps`.
        Default is 1e-5.
    maxiter : int
        The maximum number of iterations (default is 1000.)
        Iterations terminate if the number of iterations exceeds `maxiter`.
        This parameter is only needed to avoid infinite loops if iterations wander off.
        
    Returns
    -------
    x : float
        The estimate for the root.
    niter : int
        The number of iterations.
    """
    next_x = 0
    x = x0
    
    condition = 1
    niter = 0
    while condition >= eps:
        next_x = x - m*f(x) / fder(x)
        condition = np.linalg.norm(next_x - x)
        x = next_x
        if niter == maxiter:
            return next_x, maxiter
        niter += 1
    return next_x, niter
```

<!-- grade_id: cell-d2c8d42721e39f15, points: 5, grade: true, solution: false -->
```python
for m in [1, 2, 3, 4, 5, 6]:
    xx, nit = mod_newton(lambda x: (x**2 - 1)**4,
                         lambda x: 4*(x**2 - 1)**3 * 2 * x,
                         x0=2, m=m, maxiter=10000, eps=1e-9)
    assert_allclose(xx, 1.0, atol=1e-8)
    
```

(Extra task, not graded)

Use the modified Newton's iteration for $f(x) = (x^2 - 1)^2$ with $m= 1, 2, 3, 4, 5, 6$. Use a fixed initial guess $x_0$ and a fixed tolerance $\epsilon$. How many iteration is needed for convergence? Does the variation of the number of iterations with $m$ support
the expectation that the convergence is quadratic for $m$ equal to the multiplicity of the root, and linear otherwise?

```python
for m in [1, 2, 3, 4, 5, 6]:
    _, nit = mod_newton(lambda x: (x**2 - 1)**2,
                         lambda x: 4*x*(x**2 - 1),
                         x0=2, m=m, maxiter=10000, eps=1e-9)
    print(f'For m={m} and eps=1e-9 it is needed {nit} iterations')
# Не знаю, почему так, но у меня последние m-ки будто расходятся, хотя кажется, что не должны бы.
```

Output:
```
For m=1 and eps=1e-9 it is needed 31 iterations
For m=2 and eps=1e-9 it is needed 6 iterations
For m=3 and eps=1e-9 it is needed 30 iterations
For m=4 and eps=1e-9 it is needed 10000 iterations
For m=5 and eps=1e-9 it is needed 10000 iterations
For m=6 and eps=1e-9 it is needed 10000 iterations
```

# III. Newton's fractal

(Extra task, not graded)

Consider the equation

$$
x^3=1
$$ 
 
It has three solutions in the complex plane,  $x_k = \exp(i 2\pi k / 3)$ ,  $k = 0,1,2$.

The Newton's iterations converge to one of these solutions, depending on the starting point in the complex plane (to converge to a complex-valued solution, the iteration needs a complex-valued starting point).

Plot the _basins of attraction_ of these roots in the complex plane of $x$  (i.e., on the plane  $Re(x)$--$Im(x)$ ). To this end, make a series of calculations, varying the initial conditions on a grid of points. Color the grid in three colors, according to the root to which the iterations have converged.


```python

```